```python
# Initialize Otter
import otter
grader = otter.Notebook("quiz9-classes.ipynb")
```

# ⌛️ Quiz 9 - Classes

This quiz will evaluate your mastery of using classes in Python. Functions provide a way to isolate code that you want to use repeatedly, and they allow you to pass in data to the code and get data back out of the code.


## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
import pkg_resources
from subprocess import call
import sys

package_name = 'ENGR131_Util_2024'
version = '0.1.11'
package_version = f'{package_name}=={version}'

try:
    # Check if the package and version are installed
    pkg_resources.require(package_version)
    print(f'{package_version} is already installed.')
except pkg_resources.DistributionNotFound:
    # If not installed, install the package
    print(f'{package_version} not found. Installing...')
    call([sys.executable, '-m', 'pip', 'install', package_version])
except pkg_resources.VersionConflict:
    # If a different version is installed, you can choose to upgrade/downgrade
    installed_packages = {dist.key: dist.version for dist in pkg_resources.working_set}
    installed_version = installed_packages.get(package_name.lower())
    print(f'{package_name} {installed_version} is installed, but {version} is required.')
    # Optionally, upgrade or downgrade the package to the required version
    call([sys.executable, '-m', 'pip', 'install', '--upgrade', package_version])
```

```python
# DO NOT MODIFY THIS CELL

from ENGR131_Util_2024 import cell_logger, StudentInfoForm, responses, upsert_to_json_file
# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', cell_logger)
responses["assignment"] = "quiz_9"

StudentInfoForm(**responses)
```

## Question: Beam Load Calculator

### Background

In architectural engineering, calculating the load a beam can support is crucial for ensuring the safety and integrity of a structure. The total load on a beam includes both the dead load (permanent or static load, e.g., the weight of the beam itself and permanent fixtures) and live load (temporary or dynamic load, e.g., weight of people, furniture, snow).

### Objectives

1. Implement a `Beam` class to model a structural beam in a building.
2. Include methods within this class to calculate the dead load, live load, and total load the beam can support.
3. Demonstrate the use of class initialization, basic math calculations, and functions calling other functions.

### Class to Implement

implement a class `Beam` with the following properties and methods:
   - Methods:
     - `__init__`: Initializes a new `Beam` instance with the properties.
         - Properties:
           - `length` (in meters)
           - `width` (in meters)
           - `depth` (in meters)
           - `material_density` (in kg/m³, representing the density of the beam material)
           - `live_load_per_m2` (in kg, representing the expected maximum live load per square meter)
     - `volume`: Calculates and returns the volume of the beam.
       - The volume of a rectangular prism (beam) is calculated as length x width x depth.
     - `dead_load`: Calculates and returns the dead load of the beam based on its volume and material density.
       - The dead load is the volume of the beam multiplied by its material density.
     - `live_load`: Calculates and returns the total live load the beam can support, based on the beam's length and the specified live load per square meter.
       - Live load is the beam's surface area (length x width) multiplied by the live load per square meter.
     - `total_load`: Calculates and returns the total load the beam can support by summing its dead and live loads.
       - the total load is the sum of the dead and live loads.
     - `print_loads`: Prints the dead, live, and total loads of the beam in the following format:
       - "Dead Load: {dead_load} kg\nLive Load: {live_load} kg\nTotal Load: {total_load} kg" where the values are extracted from the object rounded to 2 decimal places.
       - Note: you must use a **single print command**, \n is used to create a new line in the print statement.

Instantiate a `Beam` object to the variable `beam` the following properties:

   - Beam dimensions: 8m (length) x 0.5m (width) x 0.25m (depth)
   - Material density: 2400 kg/m³ (concrete)
   - Live load per square meter: 500 kg/m²


```python
# Your Class for a Beam goes Here
...

# Instantiate a Beam object with the specified properties
...

# Print the dead, live, and total loads of the beam
# you can uncomment this line for testing
# beam.print_loads()
```

```python
grader.check("q1-beam-load")
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/Ok5XgX3N) add link to navigate to the assignment on GitHub classroom.

If you need further instructions on submitting your assignment please look at Lab 1. 

## Viewing your score

Each `.ipynb` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

We have both public and hidden tests. You will be able to see the score of both tests, but not the specific details of why the test passed or failed. 

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q1-beam-load (points: 42)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> points_ = [6, 4, 5, 4, 4, 4, 4, 6, 5]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> import json
>>> student_data = {'first_name': responses['first_name'], 'last_name': responses['last_name'], 'drexel_id': responses['drexel_id'], 'drexel_email': responses['drexel_email']}
>>> with open('student_data.json', 'w') as json_file:
...     json.dump(student_data, json_file)
>>> scorer = submit_score()
>>> question_id = 'q1_1'
>>> max_score = 6
>>> score = 0
>>> for method in ['__init__', 'volume', 'dead_load', 'live_load', 'total_load', 'print_loads']:
...     if hasattr(Beam, method):
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': dir(Beam)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert hasattr(Beam, '__init__'), '__init__ method is not implemented'
>>> assert hasattr(Beam, 'volume'), 'volume method is not implemented'
>>> assert hasattr(Beam, 'dead_load'), 'dead_load method is not implemented'
>>> assert hasattr(Beam, 'live_load'), 'live_load method is not implemented'
>>> assert hasattr(Beam, 'total_load'), 'total_load method is not implemented'
>>> assert hasattr(Beam, 'print_loads'), 'print_loads method is not implemented'
```
success_message: All methods are implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_2'
>>> max_score = 4
>>> score = 0
>>> if len(Beam.__init__.__code__.co_varnames) == 6:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': dir(beam)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert len(Beam.__init__.__code__.co_varnames) == 6, '__init__ method does not take the correct number of parameters'
```
success_message: Init takes the correct number of parameters
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> scorer = submit_score()
>>> question_id = 'q1_3'
>>> max_score = 5
>>> score = 0
>>> for method in ['length', 'width', 'depth', 'material_density', 'live_load_per_m2']:
...     if hasattr(beam, method):
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': dir(beam)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert hasattr(beam, 'length'), 'length attribute is not implemented'
>>> assert hasattr(beam, 'width'), 'width attribute is not implemented'
>>> assert hasattr(beam, 'depth'), 'depth attribute is not implemented'
>>> assert hasattr(beam, 'material_density'), 'material_density attribute is not implemented'
>>> assert hasattr(beam, 'live_load_per_m2'), 'live_load_per_m2 attribute is not implemented'
```
success_message: All properties are implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_4'
>>> max_score = 4
>>> score = 0
>>> b1 = Beam(length=9, width=2, depth=3, material_density=2400, live_load_per_m2=400)
>>> test = b1.volume() == 54
>>> statement = 'volume method is not implemented correctly'
>>> value = b1.volume()
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: volume method is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_5'
>>> max_score = 4
>>> score = 0
>>> b1 = Beam(length=9, width=2, depth=3, material_density=2400, live_load_per_m2=400)
>>> value = b1.dead_load()
>>> test = b1.dead_load() == 129600
>>> statement = 'dead load method is not implemented correctly'
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: dead load implemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_6'
>>> max_score = 4
>>> score = 0
>>> b1 = Beam(length=9, width=2, depth=3, material_density=2400, live_load_per_m2=400)
>>> value = b1.live_load()
>>> test = b1.live_load() == 7200
>>> statement = 'live load method is not implemented correctly'
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: dead load implemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_7'
>>> max_score = 4
>>> score = 0
>>> b1 = Beam(length=9, width=2, depth=3, material_density=2400, live_load_per_m2=400)
>>> value = b1.total_load()
>>> test = b1.total_load() == 136800
>>> statement = 'total load method is not implemented correctly'
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: dead load implemented correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_8'
>>> max_score = 6
>>> expected_message = 'Dead Load: 129600.00 kg\nLive Load: 7200.00 kg\nTotal Load: 136800.00 kg'
>>> with patch('builtins.print') as mock_print:
...     b1 = Beam(length=9, width=2, depth=3, material_density=2400, live_load_per_m2=400)
...     b1.print_loads()
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message:
...     score = 6
...     drexel_jupyter_logger.variable_logger_csv(f'6, {max_score}', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': f'{mock_print.call_args[0][0]}'}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
```
success_message: Input parameters is is correctly defined.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> import numpy as np
>>> scorer = submit_score()
>>> question_id = 'q1_9'
>>> max_score = 5
>>> score = 0
>>> for (method, value) in zip(['length', 'width', 'depth', 'material_density', 'live_load_per_m2'], [8, 0.5, 0.25, 2400, 500]):
...     if eval(f'beam.{method}') == value:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': f'Length: {beam.length}, \n                                Width: {beam.width}, \n                                Depth: {beam.depth}, \n                                Material Density: {beam.material_density}, \n                                Live Load per m2: {beam.live_load_per_m2}\n                                '}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> for (method, value) in zip(['length', 'width', 'depth', 'material_density', 'live_load_per_m2'], [8, 0.5, 0.25, 2400, 500]):
...     assert eval(f'beam.{method}') == value
```
success_message: Input parameters is is correctly defined.
